---
title       : Introdução
description : Neste módulo você vai aprender fazer operações matemáticas no console, criar objetos (variáveis), pedir ajuda, e os principais tipos (classes) de dados.
attachments :

--- type:NormalExercise lang:r xp:100 skills:1,3 key:bc438b12bd
## Operações matemáticas 1

O exercício é usar o R como Calculadora 

*** =instructions
1.  Some 10 + 20
2.  Divida 20 / 5
*** =hint
Basta seguir as instruções ;-).

 *** =pre_exercise_code
  ```{r}
  10 + 20
  20 / 5
  ```
  
  *** =sample_code
  ```{r}
  # Some 10 + 20
  ```
  
  *** =solution
  ```{r}
  # Some 10 + 20
  10 + 20
  # Divida 20 / 5
  20 / 5
  ```
  
  *** =sct
  ```{r}
  test_error()
  success_msg("Muito Bom! Siga para o proximo exercicio")
  ```
  
--- type:NormalExercise lang:r xp:100 skills:1,3 key:ac898f9bbb
  
## Operações matemáticas 2

O exercício é usar o R como Calculadora 

*** =instructions
1.  Multiplique 7 * 4
2.  Subtraia 70 - 19
*** =hint
Basta seguir as instruções ;-).

 *** =pre_exercise_code
  ```{r}
  7 * 4
  70 - 19
  ```
  
  *** =sample_code
  ```{r}
  # Multiplique 7 * 4
  # Insira sua resposta na proxima linha
  ```
  
  *** =solution
  ```{r}
  # Multiplique 7 * 4
  7 * 4
  # Subtraia 70 - 19
  70 - 19
  ```
  
  *** =sct
  ```{r}
  test_correct("7 * 4")
  success_msg("Muito Bom! Siga para o proximo exercicio")
  ```
