---
title       : Introdução
description : Neste módulo você vai aprender fazer operações matematica no console, criar objetos (variáveis), pedir ajuda, e os principais tipos (classes) de dados.

framework : datamind
---
## How it works

In the **editor** on the right you will have to type your code for solving the exercises. Now it just contains `3+4` and `demo("graphics")`. 

Since understanding and structuring long parts of code is not always straight-forward, R makes use of the `#` sign to add comments. Just like Twitter! Comments are not run as R-code, so they will not influence your result. 

The ouput of you R code is shown in the **console** in the lower right corner, while graphs will be shown in the upper right corner.

*** =instructions

1. Click Submit Answer and see how the console now shows you the executed R code: the solution `7` appears as the sum of `3` and `4`. In its most basic form R can thus be used as a calculator or to generate plots, but there's much more ;-).
2. Use the arrows to browse through some cool visualizations generated with R, and shown here in upper right corner. Note that you can make the graph window larger, if you would like to have a closer look.

*** =hint
Just click the Submit Answer button on the right.

*** =sample_code
```{r eval=FALSE}
# This  is the editor and the part below the white line is called the console.

# The  hashtag is used to add comments

# Show some demo graphs generated with R
demo("graphics")

# Calculate 3+4
3+4
```

*** =solution
```{r eval=FALSE}
#Just click the Submit Answer button!

# This  is the editor and the part below the white line is called the console.

# The  hashtag is used to add comments

# Calculate 3+4
3+4

# Show some demo graphs generated with R
demo("graphics")
```

*** =sct
```{r eval=FALSE}
DM.result <- TRUE
```

*** =pre_exercise_code
```{r eval=FALSE}

```
---
