---
title       : Introdução
description : Neste módulo você vai aprender fazer operações matemáticas no console, criar objetos (variáveis), pedir ajuda, e os principais tipos (classes) de dados.
attachments :

--- type:NormalExercise lang:r xp:100 skills:1,3 key:bc438b12bd
## Operaçoes matemáticas

O exercício é usar o R como Calculadora 

*** =Instruções
- some `10` mais  `20` no editor de script no lado direito depois submeta a resposta (Submit Answer).

*** =hint
Use `+` para somar.

*** =pre_exercise_code
```{r}
10 + 20
```

*** =sample_code
```{r}
# Some 10 + 20
```

*** =solution
```{r}
# Some 10 + 20
10 + 20
```

*** =sct
```{r}
test_error()
test_operator("+")
}
incorrect_msg = "O perador matemático para soma incorreto")
}

success_msg("Good job! Head over to the next exercise")
```

