---
title       : Introdução
description : Neste módulo você vai aprender fazer operações matemáticas no console, criar objetos (variáveis), pedir ajuda, e os principais tipos (classes) de dados.
attachments :

--- type:NormalExercise lang:r xp:20 skills:1,3 key:bc438b12bd
## Operaçoes matemáticas

O exercício é usar o R como Calculadora 

***=instructions
1.  Some 10 + 20
2.  Divida 20 / 5
*** =hint
Basta seguir as instruções ;-).

 *** =pre_exercise_code
  ```{r}
  10 + 20
  20 / 5
  ```
  
  *** =sample_code
  ```{r}
  # Some 10 + 20
  ```
  
  *** =solution
  ```{r}
  # Some 10 + 20
  10 + 20
  #Divida 20 / 5
  20 / 5
  ```
  
  *** =sct
  ```{r}
  test_error()
  success_msg("Muito Bom! Siga para o proximo exercicio")
  ```
